```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05.ipynb")
```

```python
import pandas as pd
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')

import seaborn as sns
import warnings
warnings.filterwarnings('ignore')
```

# Lab 5: Intro to Visualization

#### Learning objects

In this notebook, you will learn about:

- How to explore datasets
- How to prepare data to be visualizaed
- The purpose of different visualizations 
- How to create and code visualizations
- How to analyze and draw insights from visualizations

#### Helpful Data Science Resources 
Here are some resources you can check out while doing this notebook and to explore data visualization further!
- [DATA 8 Textbook](https://inferentialthinking.com/chapters/07/Visualization.html) - Visualization chapter
- [Reference Sheet for the datascience Module](http://data8.org/sp22/python-reference.html)
- [Documentation for the datascience Modules](http://data8.org/datascience/index.html)
- [Cool Data Visualizations](https://www.tableau.com/learn/articles/best-beautiful-data-visualization-examples)
- [Statistica: Find Data on Interesting Topics](https://www.statista.com/)
- [Exploratory Data Analysis](https://en.wikipedia.org/wiki/Exploratory_data_analysis)

**A note on the autograder for this lab:** The test cases in the autograder are not comprehensive -- you can create some completely incorrect graph while passing the autograder. So assure yourself that you are making the correct visualizations, we have provided a [list of reference images](https://www.econ148.org/sp23/resources/assets/supp_materials/lab05/reference_images/). Your score for this lab will still solely depend on the autograder provided. 

---
## Part 1: Explorative Data Analysis (EDA)

### Unemployment rate and NaN values 

Let's start off with loading in the dataset. We will be using unemployment rate data from Fred. 

**Question 1.1:** Load in the dataset `data/unemployment_rate.csv` and read it into a Pandas dataframe. Name it `unemployment_df`.

```python
unemployment_df = ...
unemployment_df.head()
```

```python
grader.check("q1_1")
```

**Question 1.2:** The dataset contains three columns: the overall unemployment rate, unemployment rate for male, and unemployment rate for female. See [Fed website](https://fred.stlouisfed.org/graph/?g=jXvf) and [Pandas Documentation](https://pandas.pydata.org/docs/refeence/api/pandas.DataFrame.rename.html) and change names of the columns to one of these : 'Date', 'Male', 'Female', and 'Unemployment rate'. 

```python
unemployment_df = ...
unemployment_df.head()
```

```python
grader.check("q1_2")
```

Before plotting the data, it is important to determine if there are any nan values in the dataset. 
You can use [pd.isnull](https://pandas.pydata.org/docs/reference/api/pandas.isna.html) or [pd.isna](https://pandas.pydata.org/docs/reference/api/pandas.isna.html) to check the existence of nan value. 

Take a look at documentation and use the knowledge to determine whether our dataset contains any NaN value. 

<!-- BEGIN QUESTION -->

**Question 1.3:**  Does the dataset contain any NaN values? Explain what you did to reach your conclusion. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---
## Part 2: Data Visualization

### Line Plots

A line plot is used to display data as a series of points connected by a line. It's generally used to visualize how a variable changes over time (also known as [time series data](https://www.investopedia.com/terms/t/timeseries.asp)). It has a time-related variable on the x-axis (minutes, days, months, years, etc.) and a numerical variable on the y-axis.

Let's create a line plot to see how the overall unemployment rate changes over the years.


**Question 2.1:** Make a line plot for the overall unemployment rate over the entire sample period. Label your plot properly (both axes and title).  

```python
def q2_1():
    unemployment_df.plot(..., ..., figsize=(8,5), style='g-', linewidth= 1.5)
    plt.ylabel(...)
    plt.xlabel(...)
    plt.title(...)
q2_1()
```

```python
grader.check("q2_1")
```

<!-- BEGIN QUESTION -->

**Question 2.2:** What kind of trend in unemployment rate can you find from the graph above? Please provide potential reasonings of the trend you find from the graph above. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 2.3:** Create similar line plots for the `unemployment_rate` dataset below **for each gender**. Label your plot properly (both axes and title).

Hint : Take a look at documentation of [df.plot()](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.plot.html) method.

```python
def q2_3():
    unemployment_df.plot(..., [..., ...], figsize=(8, 5), linewidth=2)\
    plt.ylabel(...)
    plt.xlabel(...)
    plt.title(...)
q2_3()
```

```python
grader.check("q2_3")
```

<!-- BEGIN QUESTION -->

**Question 2.4:** What kind of difference between two genders can you see from the plot above? What could be the possible reasons behind this difference? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Inflation and Unemployment rate

Now we will be using both Unemployment rate and Inflation rate at the same time to plot a graph. First let's create `inflation_df` using `inflation_rate.csv`. Let's import the dataframe and rename columns so we can use them to plot graph.  

**Question 2.5**: Load in the dataset `data/inflation_rate.csv` and read it into a Pandas dataframe. Name it `inflation_df`. Then rename the columns to `Date` and `Inflation rate`. 

```python
inflation_df = ...
...
inflation_df.head()
```

```python
grader.check("q2_5")
```

### Merging datasets

It makes easier to plot if we have a single dataframe to plot. There are many plotting libraries that require you to pass the dataframe of your interest as one of arguments. Threrefore, we will create a Dataframe that contains both `unemployment_rate` and `inflation_rate`. We can use `Date` column of both dataframe as a key to combine two tables. 

We want to create `merged_df` without any row containing NaN value. In other words, keep rows that only exsit in both dataframe.

**Question 2.6**: Create `merged_df` by inner merging `unemployment_df` and `inflation_df`. 

Hint: Read documentation of [pd.merge()](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.merge.html). Think about what type of merge you want to use. 

```python
merged_df = ...
merged_df.head()
```

```python
grader.check("q2_6")
```

Now, lets create a plot using the dataframe we just created. This time we will be using `plt.plot()` method to plot multiple graphs into a single plot. Feel free to read documentatioon of [plt.plot()](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.plot.html) and fill in `...` to create the plot.

**Question 2.7**: Fill in the blanks to plot two line graphs for both Inflation rate and Unemployment rate and label their names using respectively. Please set the name of x and y label to `Date` and `Percentage (%)`. Label your plot properly (both axes and title).

```python
def q2_7():
    fig, ax = plt.subplots(figsize=(8,5), dpi=100)

    ax.plot(..., ..., label=..., linewidth=2) # inflation
    ax.plot(..., ..., label=..., linewidth=2) # unemployment

    ticks = merged_df["Date"].loc[range(0,len(merged_df["Date"]), 15)].values # this will create ticks on the x-axis
    plt.xticks(ticks)
    
    ... # label your plot

    plt.legend()
q2_7()
```

```python
grader.check("q2_7")
```

### Bar Charts

Another visualization we can use to see how the time spent on each activity changes over the years is a bar chart. **A bar chart is used to visualize the distribution of a categorical variable.** The bar chart is a familiar way of visualizing categorical distributions. It displays a bar for each category where the bars are equally spaced and equally wide. The length of each bar is proportional to the frequency of the corresponding category. 

In this case, we don't have categorical variable in our both datasets. However, if we create an extra column that only takes integer value from both `inflation_rate` and `unemployment_rate` columns and plot a bar chart using the column, we can get the distribution that looks the same as the distribution of histogram. 

### Histograms

**A histogram allows you to visualize the distribution of a numerical variable.** It helps you understand how spread out the values in your table are. A histogram is a visualization of the distribution of a quantitative variable. It looks very much like a bar chart but there are some **important differences** that we will examine in this section.

Histograms follow the **area principle** and have two defining properties:

1. The bins are drawn to scale and are contiguous (though some might be empty), because the values on the horizontal axis are numerical and therefore have fixed positions on the number line.
1. The area of each bar is proportional to the number of entries in the bin.

The histogram is said to be drawn on the *density scale*. The area of each bar is equal to the percent of data values that are in the corresponding bin.

> **`.hist`**: create a histogram (specify the name of the column, plus other optional arguments)
\

**Question 2.8**: Plot histogram for inflation rate using `.hist()` method. Create `bins` array for histogram from 3 to 15 with stepsize of 1.5. Label the axes to `Inflation rate (%)` and `Counts` and name the title of the plot `Distribution of Inflation rate (%)`. 

```python
def q2_8():
    fig, ax = plt.subplots(figsize =(8, 5), dpi=100)

    bins = ...
    ax.hist(x=..., bins=..., rwidth=.75)
    ... # label your plot
q2_8()
```

```python
grader.check("q2_8")
```

**Question 2.9**: Plot histogram for Unemployment rate using .hist() method. Create `bins` array for histogram from 3 to 11 with stepsize of 0.5. Label the axes `Unemployment rate (%)` and `Counts` and name the title of the plot as `Distribution of Unemployment rate (%)`.

```python
def q2_9():
    fig, ax = plt.subplots(figsize =(8, 5), dpi=100)

    bins = ...
    ax.hist(x=..., bins=..., rwidth=.75)
    ... # label your plot
q2_9()
```

```python
grader.check("q2_9")
```

### Scatter Plots

Another type of data visualization that is really important in statistics is the scatter plot. **Scatter plots are used to visualize the relationship between two numerical variables.**

Specifically, it allows us to infer the association between two variables. Association refers to how one variable (the y-variable) changes with respect to the other variable (the x-variable). See [this chapter](https://inferentialthinking.com/chapters/15/Prediction.html) in the DATA 8 textbook for reference.

We can describe the association between two variables based on two factors:

1. Magnitude: Is the association strong or weak? If the points on the scatter plot all line up along a straight line (in any direction) or "blob", it means that the association between the variables is strong. On the other hand, if the points are all spread out, it means that the association is weak.

2. Direction (or sign): Is the association positive or negative? If the points on the scatter plot line up along an upward sloping line, it means that the association between the variable is positive. If they line up along a downward sloping line, the association is negative.

Now that we understand what scatter plots are and how to interpret them, let's create one to visualize the association between Unemployment rate and Inflation rate.

> **`.scatter`:** create a scatter plot (the first argument is the variable on the x-axis and the second argument is the variable on the y-axis)


**Question 2.10** Create a scatter plot using `Unemployment rate` for x-axis and `Inflation rate` for y-axis. Also draw a linear line of the best fit with color red using `slope` and `intercept` returned from [`np.polyfit()`](https://numpy.org/doc/stable/reference/generated/numpy.polyfit.html). 

```python
def q2_10():
    fig, ax = plt.subplots(figsize =(5, 5), dpi=100)

    x = ...
    y = ...

    ax.scatter(..., ...)

    #obtain m (slope) and b(intercept) of linear regression line
    m, b = np.polyfit(..., ..., ...)

    #use red as color for regression line (line of the best fit)
    plt.plot(..., ..., ...)

    ... # label your plot
q2_10()
```

```python
grader.check("q2_10")
```

### Philips curve

The ***Phillips Curve*** is an economic concept that shows the relationship between inflation and unemployment. It suggests that there is an inverse relationship between these two variables, meaning that as unemployment decreases, inflation increases, and vice versa. The curve is based on the idea that as the labor market tightens, employers must offer higher wages to attract workers, which increases production costs and leads to higher prices for goods and services. 

You can get more information about the Philips curve [here](https://en.wikipedia.org/wiki/Phillips_curve)



One of reasons why our regrassion line above was not downward sloping is because we have too broad time frame with small number of data points for each year. Lets categorize each data point repective to decades. You can do this by simply creating extra column with column name `Decade` and assign the decade value calculated using `Date` column. 

**Question 2.11**: Create `Decade` column in `merged_df` that store the decade. 

Note: For example, 1967 should have decade as 1960, and 2021 should have decade as 2020.   
Hint: You can get part of a string by slicing a string. We can do this on a column in a dataframe using the string method. This [tutorial](https://note.nkmk.me/en/python-pandas-str-slice/) may be helpful.

```python
merged_df["Decade"] = ... + '0' # hint: for example, '195' + '0'
merged_df.head()
```

```python
grader.check("q2_11")
```

Now we will use the [seaborn](https://seaborn.pydata.org/index.html) package to create a scatter plot using the same data above but with different colors based on the `Decade` column. Read documentation of [`sns.scatterplot()`](https://seaborn.pydata.org/generated/seaborn.scatterplot.html) to create a scatter plot with different colors based on `Deceade`. Still we want to have unemployment rate on the x-axis and inflation rate on the y-axis.

**Question 2.12**: Fill in the `...` to create a scatter plot of our interest. Also, label your plot properly (both axes and title). 

```python
def q2_12():
    fig, ax = plt.subplots(figsize=(5, 5), dpi=100)

    sns.scatterplot(x=..., y=..., data=..., hue=..., palette="bright", alpha=0.75)

    ... # label your plot
q2_12()
```

```python
grader.check("q2_12")
```

Although it is possible to use the `np.polyfit()` and `plt.plot()` to graph the regression line, we will be introducing new plot `sns.lmplot` from Seabon library. This is a powerful tool in the sense that use can create and color the scatter plot and regression lines based on the categorical column of your choice at the same time. Read the documentation of [`sns.lmplot`](https://seaborn.pydata.org/generated/seaborn.lmplot.html) for more information.

**Question 2.13**: Fill in the `...` to create a `lmplot` of our interest. Use `Unemployment rate` and `Inflation rate` from `merged_df_without_20s` to create a plot with their color based on `Decade` column. Also, label both x-axis and y-axis of the plot and title the plot as same as above. 

Note: As the decade `2020` has only three data points, please create `merged_df_without_20s` filtering out data in the decade `2020`.  

```python
def q2_13():
    merged_df_without_20s = ...

    sns.lmplot(x=..., y=..., data=..., hue=..., 
            palette="bright", ci=None, line_kws={'lw': 2}, lowess =True)

    ... # label your plot
    return merged_df_without_20s
_ = q2_13()
```

```python
grader.check("q2_13")
```

---
## Part 3. More Plots with Seaborn

### Boxplot

A boxplot is similar to a histogram as it also **visualizes the distribution of a numerical variable, but it gives you more specific statistics about the distribution**: the minimum, lower quartile (value at the bottom 25th percentile), median (value at the 50th percentile), upper quartile (value at the top 25th percentile) and maximum. [Read here](https://seaborn.pydata.org/generated/seaborn.boxplot.html) for information about box plots.


**Question 3.1**: Use [`sns.boxplot()`](https://seaborn.pydata.org/generated/seaborn.boxplot.html) on `unemployment_df` to create a box plot of three different unemployment rates: `Male`, `Unemployment rate`, and `Female`. Title the plot as `Distribution of Unemployment rates (%)` and label y-axs as `Unemployment rate (%)`

```python
def q3_1():
    sns.boxplot(...)
    ... # label your plot
q3_1()
```

```python
grader.check("q3_1")
```

<!-- BEGIN QUESTION -->

**Question 3.2**: Use the boxplot above to find *approximate* values for the following statistics for the Unemployment rate (overall, female, and male). 
- **Minimum**: ...
- **Lower quartile**: ...
- **Median**: ...
- **Upper quartile**: ...
- **Maximum**: ...

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Violin Plot

A violin plot is a combination of a histogram and boxplot. **It shows you the general distribution of the data (by creating a histogram and drawing a line to capture its general shape) as well as specific statistics (same as the ones a boxplot shows).** [Read here](https://chartio.com/learn/charts/violin-plot-complete-guide/) for information about violin plots.

**Question 3.3**: Use [`sns.violinplot()`](https://seaborn.pydata.org/generated/seaborn.violinplot.html) on `merged_df` to create a violin plot of `Unemploymnet rate` and `Inflation rate`.  Set the palette argument of the violin plot to `prism`. Also, title the plot as `Distribution of Inflation rate & Unemployment rate` and label y-axis as `Percentage (%)`.

```python
def q3_3():
    sns.violinplot(...)
    plt.title(...)
    plt.ylabel(...)
q3_3()
```

```python
grader.check("q3_3")
```

<!-- BEGIN QUESTION -->

**Question 3.4**: What did you notice from the violin plot above? (Distribution, skewness, etc.) Please give the reasoning to your answer.   

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---
## Conclusion

Over the course of this notebook, you were introduced to basic manipulation and analysis of data tables that is frequently necessary in order to create visualizations. Now you've learned about different types of visualizations, practiced creating various plots, and know what visualizations are best for each goal. Feel free to reference this notebook for your Survey Analysis project.

**Congratulations!** You have finished the notebook!

---
## Feedback

**Question 4:** Please fill out this short [feedback form](https://forms.gle/4PgR9FmW37w36dR48) to let us know your thoughts about this lab! We really appreciate your opinions and feedback! At the end of the Google form, you should see a codeword. Assign the codeword to the variable `codeword` below. 

```python
codeword = ...
```

```python
grader.check("q4")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: None)
```
>>> unemployment_df.shape == (826, 4)
True
```


Tests for q1_2 (points: None)
```
>>> unemployment_df.shape == (826, 4)
True
```
```
>>> np.all(unemployment_df.columns == ['Date', 'Male', 'Female', 'Unemployment rate'])
True
```


Tests for q2_1 (points: None)
```
>>> q2_1()
>>> ax = plt.gca()
>>> plt.close()
>>> data = ax.lines[0].get_xydata().T
>>> 
>>> assert data.shape == (2, 826)
>>> assert np.isclose(np.sum(data[0]), 340725.0)
>>> assert np.isclose(np.sum(data[1]), 4856.3)
>>> assert len(ax.get_title()) > 0
>>> assert len(ax.get_xlabel()) > 0
>>> assert len(ax.get_ylabel()) > 0
```


Tests for q2_10 (points: None)
```
>>> q2_10()
>>> ax = plt.gca()
>>> plt.close()
>>> data = ax.lines[0].get_xydata().T
>>> 
>>> assert data.shape == (2, 68)
>>> assert np.isclose(np.sum(data[0]), 397.6)
>>> assert np.isclose(np.sum(data[1]), 245.43425)
>>> assert len(ax.get_title()) > 0
>>> assert len(ax.get_xlabel()) > 0
>>> assert len(ax.get_ylabel()) > 0
```


Tests for q2_11 (points: None)
```
>>> len(merged_df[merged_df["Decade"] == '1950']) == 5
True
```
```
>>> len(merged_df[merged_df["Decade"] == '2020']) == 3
True
```


Tests for q2_12 (points: None)
```
>>> q2_12()
>>> ax = plt.gca()
>>> plt.close()
>>> cs = ax.collections[0]
>>> data = cs.get_offsets().T
>>> assert data.shape == (2, 68)
>>> assert np.isclose(np.sum(data[0]), 397.6)
>>> assert np.isclose(np.sum(data[1]), 245.43425)
>>> assert len(ax.get_title()) > 0
>>> assert len(ax.get_xlabel()) > 0
>>> assert len(ax.get_ylabel()) > 0
```


Tests for q2_13 (points: None)
```
>>> df_output = q2_13()
>>> assert df_output.shape == (65, 6)
>>> 
>>> ax = plt.gca()
>>> plt.close()
>>> cs = ax.collections[0]
>>> data = cs.get_offsets().T
>>> assert data.shape[0] > 0
>>> assert data.shape[1] > 0
>>> assert len(ax.get_title()) > 0
>>> assert len(ax.get_xlabel()) > 0
>>> assert len(ax.get_ylabel()) > 0
```


Tests for q2_3 (points: None)
```
>>> q2_3()
>>> ax = plt.gca()
>>> plt.close()
>>> data = ax.lines[0].get_xydata().T
>>> 
>>> assert data.shape == (2, 826)
>>> assert np.isclose(np.sum(data[0]), 340725.0)
>>> assert np.sum(data[1]) > 0
>>> assert len(ax.get_title()) > 0
>>> assert len(ax.get_xlabel()) > 0
>>> assert len(ax.get_ylabel()) > 0
```


Tests for q2_5 (points: None)
```
>>> inflation_df.shape == (75, 2)
True
```
```
>>> np.all(inflation_df.columns == ['Date', 'Inflation rate'])
True
```


Tests for q2_6 (points: None)
```
>>> merged_df.shape == (68, 5)
True
```
```
>>> np.allclose(merged_df[merged_df["Date"] == "1959/1/1"]["Unemployment rate"], 6.0)
True
```
```
>>> np.allclose(merged_df[merged_df["Date"] == "1959/1/1"]["Inflation rate"], 0.93199)
True
```


Tests for q2_7 (points: None)
```
>>> q2_7()
>>> ax = plt.gca()
>>> plt.close()
>>> data = ax.lines[0].get_xydata().T
>>> 
>>> assert data.shape == (2, 68)
>>> assert np.isclose(np.sum(data[0]), 2278.0)
>>> assert np.isclose(np.sum(data[1]), 245.43425000000002)
>>> assert len(ax.get_title()) > 0
>>> assert ax.get_xlabel() == 'Date'
>>> assert ax.get_ylabel() == 'Percentage (%)'
```


Tests for q2_8 (points: None)
```
>>> q2_8()
>>> ax = plt.gca()
>>> plt.close()
>>> p = ax.patches
>>> assert p[0].get_xy() == (3.1875, 0.0)
>>> assert p[0].get_height() == 16
>>> assert ax.get_title() == 'Distribution of Inflation rate (%)'
>>> assert ax.get_xlabel() == 'Inflation rate (%)'
>>> assert ax.get_ylabel() == 'Counts'
```


Tests for q2_9 (points: None)
```
>>> q2_9()
>>> ax = plt.gca()
>>> plt.close()
>>> p = ax.patches
>>> assert p[0].get_xy() == (3.0625, 0.0)
>>> assert p[0].get_height() == 1
>>> assert ax.get_title() == 'Distribution of Unemployment rate (%)'
>>> assert ax.get_xlabel() == 'Unemployment rate (%)'
>>> assert ax.get_ylabel() == 'Counts'
```


Tests for q3_1 (points: None)
```
>>> q3_1()
>>> ax = plt.gca()
>>> plt.close()
>>> assert len(ax.lines) > 0
>>> assert ax.get_title() == 'Distribution of Unemployment rates (%)'
>>> assert ax.get_ylabel() == 'Unemployment rate (%)'
```


Tests for q3_3 (points: None)
```
>>> q3_3()
>>> ax = plt.gca()
>>> plt.close()
>>> assert len(ax.lines) > 0
>>> assert ax.get_title() == 'Distribution of Inflation rate & Unemployment rate'
>>> assert ax.get_ylabel() == 'Percentage (%)'
```


Tests for q4 (points: None)
```
>>> sum([ord(c) for c in codeword])^0xec148 == 967586
True
```
