```python
# Initialize Otter
import otter
grader = otter.Notebook("lab00.ipynb")
```

```python
import math
import numpy as np
import pandas as pd
```

# Lab 0: Intro to Notebook

Welcome to the first lab for this course! We'll go over some Python and NumPy exercises. For most of you, this will be a refresher on how to use Jupyter notebook and how to write code in Python. We will also go over some mathematics that are helpful for this course. 

#### Learning Objectives
1. Review basic Python syntax
2. Review baisc numpy syntax and useful methods
3. Introduce basic multivariate calculus and linear algebra and their implementation in Python

----

## Section 1: Python  <a id='python'></a>

Python is the main programming language we'll use in the course. Work through the following cells and make sure you understand what is happening in each.

Feel free to review one or more of the following materials as a refresher.

- **[Python Tutorial](https://docs.python.org/3/tutorial/index.html)**: Introduction to Python from the creators of Python
- **[Composing Programs Chapter 1](http://composingprograms.com/pages/11-getting-started.html)**: This is more of a introduction to programming with Python.
- **[Advanced Crash Course](http://cs231n.github.io/python-numpy-tutorial/)**: A fast crash course which assumes some programming background.
- **[Jupyter Notebook Basics](https://jupyter-notebook.readthedocs.io/en/stable/examples/Notebook/Notebook%20Basics.html)**: A guide to navigating Jupyter Notebooks, including different parts of the user interface and different ways that cells can be used.

### Mathematical Expressions

In this section, we'll go over some basic mathematical expressions in Python.

```python
# In Python, the ** operator performs exponentiation.
math.sqrt(math.e ** (-math.pi + 8))
```

**Question 1.1:** Multiply twenty by the sum of ten raised to the power of two and one. Then, add three. 

```python
q1_1 = ...
q1_1
```

```python
grader.check("q1_1")
```

### Output and Printing

<!-- BEGIN QUESTION -->

**Question 1.2:** One at a time, uncomment and run each of the three lines of code below. Without modifying the text itself, change the second line of code so that it prints without errors. 

```python
# print('Hello' + ',', 'world!')
# 'Why didn't this line print?'
# "Hello, cell" + " output!"
```

<!-- END QUESTION -->

### Strings ([Reference](https://developers.google.com/edu/python/strings))

Run the code in the two cell blocks below. Notice the errors that occur when you uncomment line 5 in the first block and line 2 in the second block.

```python
s = 'hi'
print(s[1])         ## i
print(len(s))       ## 2
print(s + ' there')
#s[1] = 'a' # why doesn't this work?
```

```python
avos_num = 6.022 * 10**23
#text = "The value of Euler's number is " + avos_num      ## NO, does not work
text = "The value of Euler's number is "  + str(avos_num)  ## yes
print(text)
```

**Question 1.3:** Write a line of code that will print the statement "One mole contains 6.022e23 parts" without explicitly typing out the number. (Note: in this case, it's okay if your code prints more than three decimal places.)

```python
...
```

### Lists ([Reference](https://docs.python.org/3.5/tutorial/introduction.html#lists))

```python
maybe_primes = [1, 2, 3, 5, 7, 13, 17] 
maybe_primes
```

Practice subsetting lists by running the code in the cell block below.

```python
print(maybe_primes[0])
print(maybe_primes[-1])
print(maybe_primes[-3:])
print(maybe_primes[:])
print(maybe_primes + [19, 23, 29, 31, 37])
```

**Question 1.4:** Insert 11 as the sixth element to the list `maybe_primes` and store the result in `more_primes` (such that `more_primes` = [1, 2, 3, 5, 7, 11, 13, 17]).

```python
more_primes = maybe_primes.copy() # make a copy of maybe_primes
...
more_primes
```

```python
grader.check("q1_4")
```

**Question 1.5:** Since 1 doesn't have two positive factors, remove 1 as the first element in `more_primes` and store the result in `primes` (such that `primes` = [2, 3, 5, 7, 11, 13, 17]).

```python
primes = more_primes.copy()
...
primes
```

```python
grader.check("q1_5")
```

### Dictionaries ([Reference](https://realpython.com/python-dicts/))

Practice building and modifying dictionaries by running the lines of code below.

```python
tel = {'duncan': 7214, 'julia': 5414, 'akshay': 1685, 'christian': 9522}

# insert a new key value pair
tel['richard'] = 3698 

# delete an existing key value pair (two ways)
del tel['julia']
tel.pop('duncan')

# get a value by key
print(tel['akshay'])
```

```python
tel
```

**Question 1.6:** Write code that remove the record for "christian" from the dictionary `tel2` and add a new record for "eric" with value 1480. 

```python
tel2 = {'duncan': 7214, 'julia': 5414, 'akshay': 1685, 'christian': 9522}
...
```

```python
grader.check("q1_6")
```

### For Loops

A for loop is used for iterating over a sequence (that is either a list, a tuple, a dictionary, a set, or a string). For example, if we want to show a message 100 times, then we can use a loop. It's just a simple example; you can achieve much more with loops.

**Question 1.7:** Write a for loop that computes all multiples of 14, 48, and 148 less than 1000, and stores this list of numbers in `multiples`. 

```python
multiples = []
for ...:
    if ...:
        multiples.append(...)
    ...
print(multiples)
```

```python
grader.check("q1_7")
```

### Defining Functions

Here's an example of a function:

```python
def add2(x):
    """This docstring explains what this function does: it adds 2 to a number."""
    return x + 2
```

#### Getting Help

```python
help(add2)
```

**Question 1.8:** Fill in the ellipses in the function `print_extension`, which takes in a dictionary `phonebook` and string `name` and returns a string stating the four-digit value associated with the key `name`. For example:

    >>> print_extension(tel, 'richard')
    'The extension for richard is 3698.'

Note: You don't need to handle cases where `name` doesn't exist in the `phonebook`. 

```python
def print_extension(phonebook, name):
    value = ... # From the dictionary `phonebook,` extract the value associated with key `name`
    return ... # Return the string "The extension for [name] is [value].", where [name] and [value] are the corresponding values
print_extension(tel, 'akshay') # this should print "The extension for akshay is 1685."
```

```python
grader.check("q1_8")
```

**Question 1.9:** Write a function `squares_dict` that takes in a list of numbers `l` and creates a dictionary that records the the square of each number.

    >>> print(squares_dict([0,1,2,4]))
    {0: 0, 1: 1, 2: 4, 4: 16}
    
It might be helpful to [review dictionaries in Python](https://docs.python.org/3/tutorial/datastructures.html#dictionaries).

```python
def squares_dict(l):
    squares = {} # initialize dictionary to hold the squared numbers

    for number in ...:
        ...

    return squares
print(squares_dict([0,1,2,4]))
```

```python
grader.check("q1_9")
```

----

## Section 2: NumPy <a id='NumPy'></a>

The `NumPy` library lets us do fast, simple computing with numbers in Python. 

If the following section is a little difficult, we've provided some resources for review:

* [DS100 Numpy Review](http://ds100.org/fa17/assets/notebooks/numpy/Numpy_Review.html)
* [Condensed Numpy Review](http://cs231n.github.io/python-numpy-tutorial/#numpy)
* [The Official Numpy Tutorial](https://numpy.org/doc/stable/user/quickstart.html)

<br>

**Jupyter pro-tip**: Pull up the docs for any function in Jupyter by running a cell with
the function name and a `?` at the end:

```python
np.arange?
```

**Another Jupyter pro-tip**: Pull up the docs for any function in Jupyter by typing the function
name, then `<Shift>-<Tab>` on your keyboard. Super convenient when you forget the order
of the arguments to a function. You can press `<Tab>` multiple tabs to expand the docs. Pressing `<Tab>` does an autocomplete, but doesn’t pull up docs.


Try it on the function below:

```python
np.linspace # what parameters will this function accept? What type do they need to be?
```

In the following questions, we will use numpy to answer some questions about the California Independent System Operator (CAISO) Daily Renewables Watch. This data records the hourly production of various renewable resources as well as a hourly breakdown of total production of resource type, both in megawatts (MW). For an idea of how the dataset is formatted, click [here](http://content.caiso.com/green/renewrpt/20180617_DailyRenewablesWatch.txt). 

The following line loads the data for Question 2.1.  Note that it uses Pandas -- we'll learn more about that next week!  But the output is a numpy array.  

```python
data1 = pd.read_csv("data/q2_1.csv").values 
data1
```

**Question 2.1:** Use numpy operations to output an array containing the hourly values of solar photovoltaic and solar thermal production. 

Hint: You may want to review how to slice a numpy array. 

```python
q2_1 = ...
q2_1
```

```python
grader.check("q2_1")
```

Now, we will import another dataset that contains all the data for renewable resources generation in April 2017. 

```python
data2 = pd.read_csv("data/q2_2.csv").values
data2
```

**Question 2.2:** Use numpy operations to answer the following question: How much more wind than geothermal electricity (in WM) was generated in April 2017 in total? Note, each entry in the array represents energy produced for a given *hour*.

Hint: You may want to use [np.sum](https://numpy.org/doc/stable/reference/generated/numpy.sum.html) or [sum](https://docs.python.org/3/library/functions.html#sum). 

```python
q2_2 = np.sum(data2[:, 6] - data2[:, 2])
q2_2
```

```python
grader.check("q2_2")
```

----

## Optional: Multivariable Calculus <a id='MultiCalc'></a>

The following section will ask you to recall your knowledge of multivariable calculus. Although the class is not a hard prerequisite, we will use fundamental concepts, so the following problems should familiarize you to these topics if you haven't taken it yet.

(Note that the math is written in a simplified version of LaTeX.  Writing out your answer involves putting exponents on other terms.  See how it works in the markdown source for this expression: $x^{2y}$

----

<!-- BEGIN QUESTION -->

**Question 3.1:** Simplify the following expression into the form 
$$ \ln(a) + f(x) $$
where $a$ is a constant and $f(x)$ is a polynomial function of $x$ (i.e. no natural logarithm):

$$
\ln \left(0.44e^{-3  x} e^{\cos{x}} \right) = \cdots
$$

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3.2**: Suppose we have the following scalar-valued function on $x$ and $y$:

$$ f(x, y) = \frac{1}{\cos^2{x}} - 4^y + \sqrt{x^3 + 3y}  $$

Compute the partial derivative with respect to $x$.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3.3:** Now compute the partial derivative of $f(x,y)$ with respect to $y$:

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Optional: Linear Algebra <a id='LinAlg'></a>

We'll continue the math review with a short (optional) section on linear algebra. Luckily, we have NumPy at our disposal to perform matrix operations.

In this question, we'll ask you to use your linear algebra knowledge to fill in NumPy matrices. To conduct matrix multiplication in NumPy, you should write code like the following:

```python
# A matrix in NumPy is simply a 2-dimensional NumPy array
matA = np.array([
    [1, 2, 3],
    [4, 5, 6],
])

matB = np.array([
    [10, 11],
    [12, 13],
    [14, 15],
])

# The notation B @ v means: compute the matrix multiplication Bv
matA @ matB
```

You can also use the same syntax to do matrix-vector multiplication or vector dot products. Handy!

```python
matA = np.array([
    [1, 2, 3],
    [4, 5, 6],
])

# A vector in NumPy is simply a 1-dimensional NumPy array
some_vec = np.array([ 10, 12, 14, ])

another_vec = np.array([ 10, 20, 30 ])

print(matA @ some_vec)
print(some_vec @ another_vec)
```

Sofia, Pete, and Rowan are shopping for fruit at Berkeley Bowl. Berkeley Bowl, true to its name, only sells fruit bowls. A fruit bowl contains some fruit and the price of a fruit bowl is the total price of all of its individual fruit.

Berkeley Bowl has bananas for \\$2.00, plums for \\$3.00, and dragon fruit for \\$7.00 (expensive!). These prices can be written in a vector:

$$
\vec{v} = \begin{bmatrix}
     2 \\
     3 \\
     7 \\
\end{bmatrix}
$$

Berkeley Bowl sells the following fruit bowls:

1. 2 of each fruit
2. 6 bananas and 2 plums
3. 5 dragon fruit




**Question 4.1:** Create a 2-dimensional numpy array encoding the matrix $B$ such that the matrix-vector multiplication

$$
B\vec{v}
$$

evaluates to a length 3 column vector containing the price of each fruit bowl. The first entry of the result should be the cost of fruit bowl #1, the second entry the cost of fruit bowl #2, etc.

```python
v = np.array([2,3,7])

# Fill in B
B = np.array([
    ...
])

# The notation B @ v means: compute the matrix multiplication Bv
B @ v
```

```python
grader.check("q4_1")
```

**Question 4.2:** Sofia, Pete, and Rowan make the following purchases:

- Sofia buys 6 fruit bowl #2s.
- Pete buys 1 of each fruit bowl.
- Rowan buys 3 fruit bowl #1s and 1 fruit bowl #3.


Create a matrix $A$ such that the matrix expression

$$
AB\vec{v}
$$

evaluates to a length 3 column vector containing how much each of them spent. The first entry of the result should be the total amount spent by Sofia, the second entry the amount sent by Pete, and the third entry the amount spent by Rowan.

```python
# Fill in A
A = ...

A @ B @ v
```

```python
grader.check("q4_2")
```

**Question 4.3:** Let's suppose Berkeley Bowl changes their fruit prices, but you don't know what they changed their prices to. Sofia, Pete, and Rowan buy the same quantity of fruit baskets and the number of fruit in each basket is the same, but now they each spent these amounts, respectively:

$$
\vec{x} = \begin{bmatrix}
    402 \\
    136 \\
    147 \\
\end{bmatrix}
$$

Use `np.linalg.inv` or `np.linalg.solve` to compute the new prices for the individual fruits.

```python
new_v = ...
new_v
```

```python
grader.check("q4_3")
```

----

**Congratulations**, you're done with the first lab in Econ 148!

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1_1 (points: None)
```
>>> np.isclose(q1_1, 2023)
True
```


Tests for q1_4 (points: None)
```
>>> more_primes == [1, 2, 3, 5, 7, 11, 13, 17]
True
```


Tests for q1_5 (points: None)
```
>>> primes == [2, 3, 5, 7, 11, 13, 17]
True
```


Tests for q1_6 (points: None)
```
>>> len(tel2) == 4
True
```
```
>>> not 'christian' in tel2
True
```
```
>>> assert 'eric' in tel2
>>> np.isclose(tel2['eric'], 1480)
True
```


Tests for q1_7 (points: None)
```
>>> len(multiples) == 96
True
```
```
>>> np.sum(multiples) == 47964
True
```


Tests for q1_8 (points: None)
```
>>> callable(print_extension)
True
```
```
>>> tel_test = {'alice': 1, 'bob': 2, 'chris': 3}
>>> assert print_extension(tel_test, 'alice') == 'The extension for alice is 1.'
>>> assert print_extension(tel_test, 'chris') == 'The extension for chris is 3.'
```


Tests for q1_9 (points: None)
```
>>> callable(squares_dict)
True
```
```
>>> squares_test_1 = [0, 1, 2, 4]
>>> squares_test_1_out = squares_dict(squares_test_1)
>>> for e in squares_test_1:
...     assert e in squares_test_1_out and squares_test_1_out[e] == e ** 2
```
```
>>> squares_test_2 = [-1, 1, 30, 2]
>>> squares_test_2_out = squares_dict(squares_test_2)
>>> for e in squares_test_2:
...     assert e in squares_test_2_out and squares_test_2_out[e] == e ** 2
```


Tests for q2_1 (points: None)
```
>>> q2_1.shape == (24, 2)
True
```
```
>>> np.isclose(np.sum(q2_1[:, 0]) / np.sum(q2_1[:, 1]), 22.97100987575661)
True
```


Tests for q2_2 (points: None)
```
>>> np.isclose(q2_2^0xec148, 379797)
True
```


Tests for q4_1 (points: None)
```
>>> B.shape == (3, 3)
True
```
```
>>> np.allclose(B, np.array([[2, 2, 2], [6, 2, 0], [0, 0, 5]]))
True
```


Tests for q4_2 (points: None)
```
>>> A.shape == (3, 3)
True
```
```
>>> np.allclose(A, np.array([[0, 6, 0], [1, 1, 1], [3, 0, 1]]))
True
```


Tests for q4_3 (points: None)
```
>>> len(new_v) == 3
True
```
```
>>> np.allclose(new_v, np.array([10, 3.5, 6]))
True
```
